---
title meta  : บทที่ 4
title       : Exploratory Data Analysis on Two Variable
description : Insert the chapter description here
--- type:NormalExercise lang:r xp:100 skills:1 key:53f1859947
## c


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
