---
title meta  : บทที่ 3
title       : Exploratory Data Analysis on One Variable - Part2
description : Insert the chapter description here
--- type:NormalExercise lang:r xp:100 skills:1 key:807fcae9f0
## b


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
